---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`
You should do this.

`@hint`


`@pre_exercise_code`
```{r}
# What do you want
```

`@sample_code`
```{r}
1 + 3 = 4
```

`@solution`
```{r}
4
```

`@sct`
```{r}
check_equal(4)
```
